<a href="https://colab.research.google.com/github/FelixMburu/CODE-CHALLENGES-AND-CHECKPOINTS/blob/main/DATA_SERIALIZATION.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Data Serialization Formats Checkpoint

This checkpoint is designed to test your understanding of the content from the Data Serialization Formats Cumulative Lab.

Specifically, this will cover:

* Reading serialized CSV data from a file into a Python object
* Extracting information from nested data structures

## Your Task: Extract Frequency Counts for Education Levels from a CSV

### Data Understanding

In this repository under the file path `salaries.csv` there is a CSV data file containing salary and demographic information. When loaded into Python as a list of dictionaries, each dictionary looks something like this:

```
{
  'Age':          '39',
  'Education':    'E - Bachelors',
  'Occupation':   'Adm-clerical',
  'Relationship': 'Not-in-family',
  'Race':         'White',
  'Sex':          'Male',
  'Target':       '<=50K'
 }
```

Most of this information is irrelevant for the current task; the one piece that you need to focus on is the `Education` key-value pair.

Your task is to create a **frequency table** where the various education levels (values associated with the `Education` keys) are encoded as keys, and the frequencies of those education levels are encoded as values.

### Requirements

#### 1. Import the Relevant Module

#### 2. Open the File and Create a List of Records

#### 3. Identify the Unique Education Levels

#### 4. Create a Frequency Table of Education Levels

#### 5. Identify the Most Common Education Level

## 1. Import the Relevant Module

In the cell below, import the module used for working with CSV data in Python:

<!-- grade_id: cell-c147f754df1614f5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import csv
```

<!-- grade_id: cell-e84ab604e6027b75, points: 0.5, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

## 2. Open the File and Create a List of Records

The file path is `salaries.csv`.

Make sure you follow these steps with the specified variable names in order to pass all tests:

1. Create a file object `salary_data_file` by opening the file with that path
2. Instantiate a `DictReader` ([documentation here](https://docs.python.org/3/library/csv.html#csv.DictReader)) using that file object
3. Cast the `DictReader` to a `list` and assign the result to `salary_data`
4. Close the `salary_data_file`

```python
# Replace None with appropriate code

# Open the file
# Instantiate a DictReader and create salary_data
salary_data_file = '/content/salaries.csv'
with open(salary_data_file) as f:
  salary_data = list(csv.DictReader(f))

# Close salary_data_file


# Visually inspecting the first few records
for record in salary_data[:5]:
    print(record)

```

Output:
```
{'Age': '39', 'Education': 'E - Bachelors', 'Occupation': 'Adm-clerical', 'Relationship': 'Not-in-family', 'Race': 'White', 'Sex': 'Male', 'Target': '<=50K'}
{'Age': '50', 'Education': 'E - Bachelors', 'Occupation': 'Exec-managerial', 'Relationship': 'Husband', 'Race': 'White', 'Sex': 'Male', 'Target': '<=50K'}
{'Age': '38', 'Education': 'B - HS Diploma', 'Occupation': 'Handlers-cleaners', 'Relationship': 'Not-in-family', 'Race': 'White', 'Sex': 'Male', 'Target': '<=50K'}
{'Age': '53', 'Education': 'A - No HS Diploma', 'Occupation': 'Handlers-cleaners', 'Relationship': 'Husband', 'Race': 'Black', 'Sex': 'Male', 'Target': '<=50K'}
{'Age': '28', 'Education': 'E - Bachelors', 'Occupation': 'Prof-specialty', 'Relationship': 'Wife', 'Race': 'Black', 'Sex': 'Female', 'Target': '<=50K'}
```

<!-- grade_id: cell-e95e687ed59c2de2, points: 1, grade: true, solution: false -->
```python
# Checking salary_data_file
assert type(salary_data_file) != None

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-15126aba4a36> in <cell line: 2>()
      1 # Checking salary_data_file
----> 2 type(salary_data_file)
      3 
      4 # PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
      5 # THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

NameError: name 'salary_data_file' is not defined
```

<!-- grade_id: cell-1a44994220dfeb56, points: 0.5, grade: true, solution: false -->
```python
# Checking salary_data
assert type(salary_data) == list

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

## 3. Identify the Unique Education Levels

Create a list `unique_education_levels` that contains all unique values associated with the `Education` key in these records, in alphabetical order.

***Hint:*** You'll need to loop over all records (dictionaries) in `salary_data` and find the value associated with the `Education` key for each

***Hint:*** The `.sort` list method or `sorted` built-in function can be used to sort strings into alphabetical order. Note that `.sort` modifies the list in place and returns `None`, whereas `sorted` does not modify the list in place but returns a sorted version.

<!-- grade_id: cell-1515986a288add22, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code (adding more lines as needed)

# initialize an empty list
unique_education_levels = []

# YOUR CODE HERE

for record in salary_data: #looping over each record in salary_data
  education_level = record.get('Education') # get value assocaited with Education key
  if education_level not in unique_education_levels:
    unique_education_levels.append(education_level) #append the unique education level

# sort the education levels alphabetically
unique_education_levels.sort()

print("Unique Education Levels:")
print(unique_education_levels)
```

Output:
```
Unique Education Levels:
['A - No HS Diploma', 'B - HS Diploma', 'C - Some College', 'D - Associates', 'E - Bachelors', 'F - Graduate Degree']
```

<!-- grade_id: cell-bd814f2a81f3ac73, points: 1, grade: true, solution: false -->
```python
# Checking unique_education_levels
assert type(unique_education_levels) == list
assert len(unique_education_levels) == 6

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

## 4. Create a Frequency Table of Education Levels

Create a dictionary `education_level_frequencies` where the keys are the unique education levels found above, and the values are the number of times that the education level appeared in the full dataset.

For example, the key `A - No HS Diploma` should have the associated value `4253`, since that education level appears 4,253 times in the dataset.

<!-- grade_id: cell-fa6ba75af40e3f7b, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code (add more lines as needed)

# initialize empty dictionary
education_level_frequencies = {}

# count the frequencies for each education level
for level  in education_level:
  if level in education_level_frequencies:
    education_level_frequencies[level]+=1
  else:
    education_level_frequencies[level]= 1

# Testing out your code
print("The most common education level appears", max(education_level_frequencies.values()), "times")
print("The least common education level appears", min(education_level_frequencies.values()), "times")
```

Output:
```
The most common education level appears 3 times
The least common education level appears 1 times
```

<!-- grade_id: cell-a0cdae6cd59456db, points: 1, grade: true, solution: false -->
```python
# Checking education_level_frequencies

# Should be a dictionary overall
assert type(education_level_frequencies) == dict

x = list(education_level_frequencies.keys())
height = list(education_level_frequencies.values())

# Should have string keys
assert type(x[0]) == str
# Should have integer values
assert type(height[0]) == int

# This plotting code should work
import matplotlib.pyplot as plt
fig, ax = plt.subplots(figsize=(12, 5))
ax.bar(x, height)
ax.tick_params(axis='x', labelrotation=45)
ax.set_title("Distribution of Education Levels")
ax.set_ylabel("Count");

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```

## 5. Identify the Most Common Education Level

Based on the above graph, which education level is most common in this dataset?

Set the value of the variable `most_common` to the string value of that education level. You can just type in the answer rather than finding this with code, but make sure that the test cell passes — it checks that your answer is one of the valid answers (hopefully helping you avoid a spelling mistake).

```python
# Replace None with appropriate code
most_common = max(education_level_frequencies, key=education_level_frequencies.get)

# Print the result
print(f"The key with the maximum value is '{most_common}' with a value of {education_level_frequencies[most_common]}")

```

Output:
```
The key with the maximum value is ' ' with a value of 3
```

<!-- grade_id: cell-1d083772aa40150e, points: 1, grade: true, solution: false -->
```python
assert type(most_common) == str
assert most_common in [
    'A - No HS Diploma',
    'B - HS Diploma',
    'C - Some College',
    'D - Associates',
    'E - Bachelors',
    'F - Graduate Degree'
]

# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS
```